# PS03 (v0.1)

DS256, Gettysburg College

Prof Eatai Roth

Due: Sep 26, 2025 5p


The questions are in no particular order in terms of difficulty, so if you feel stuck, skip the question and come back to it.

You'll need to run the cell below to install the otter-grader package. You'll only need to do this once, but there's no harm in running the cell multiple times.

**Your Name**: ...

**Collaborators**: ...

By submitting this work, you attest that: 
 - it reflects your own effort and understanding of the material
 - you did not ask for help from students outside the course, other than PLAs
 - you did not use a generative AI to produce any code used below.

```python
# Initialize Otter
import otter
grader = otter.Notebook("ps03.ipynb")
```

**1.** Write a function ```factorize()``` that takes as input an integer and returns as output a list of factors of that number. If the number is prime, also print "That's prime."

The function ```range(start, stop, step)``` creates a sequence of numbers as a list.

```python
def factorize(num):
...

```

```python
# Is it doing what you expect it to do.
print(factorize(1)) # should be [1]
print(factorize(57)) # should be [1, 3, 19, 57]
print(factorize(59)) # should print "That's prime." and return [1, 59]
```

```python
grader.check("q1")
```

**2.**

Write the function ```check_cyclic_equality()``` that takes two lists and checks if there is a rotation that makes them same.

A single rotation/cycling is defined as taking the zero-th element of the list and moving it to the end. For example, [‘dog’, 7, ‘chicken’, ‘rabbit’, 9] cycled by 2 would be [‘chicken’, ‘rabbit’, 9, ‘dog’, 7].

If the lists are cyclically equal, return the smallest number the second list has to be cycled to equal the first list. If not, return False.

```python
def check_cyclic_equality(lista, listb):
...

```

```python
# Lists are cyclically equal with 2 rotations
assert check_cyclic_equality([1, 2, 3], [3, 1, 2]) == 1

# Lists are not cyclically equal
assert check_cyclic_equality([1, 2, 3], [4, 5, 6]) == False
```

```python
grader.check("q2")
```

**3.** Write the function ```extract_vowels()``` that takes in a string of text and returns how many of each vowel (a,e,i,o,u) and the total number of vowels in the string.

*Hints*: 
- You can consider a string a list of characters and loop over the characters.
- 'A' and 'a' are different characters. It might be helpful (but not necessary) to convert all letters to lowercase first using ```.lower()``` (demo below).

```python
hi = 'Hello!'
for letter in hi:
    print(letter)
    
'Gettysburg College'.lower()
```

```python
...

# Assert is a kind of test statement. If the statement is True (correct), assert does nothing.
assert extract_vowels("Gettysburg College") == [0, 3, 0, 1, 1, 5]

```

```python
grader.check("q3")
```

**4.** 

Write a function, ```find_filetype()``` which will take as input a list of filenames and a file extension starting with a '.' (e.g. '.pdf', '.xlsx', '.ipynb') and return the sublist of only those files with the specified filetype.

```python
...

file_list = ['hi.pdf', 'hello.doc', 'ps01.ipynb', 'goodbye.pdf', 'test.ipynb']

find_filetype(file_list, '.ipynb')
```

```python
grader.check("q4")
```

**5.** Write the function ```remove_duplicates()``` that takes as an input a list and removes any duplicates from the list, returning a list of only unique values in the order they appear in the original list.

*Hint*: As you build your list of unique values, you'll be checking whether new values are already in that list.

```python
...

assert remove_duplicates([9, 1, 1, 'Emergency']) == [9, 1, 'Emergency']
assert remove_duplicates('Hello there!') == ['H', 'e', 'l', 'o', ' ', 't', 'h', 'r', '!']
```

```python
grader.check("q5")
```

**6. (Optional challenge)** Write the function 'convert_to_numeric()' that takes as input a list of prices. Some prices are listed as numbers, others as strings. Strings may appear with or without a decimal (but guaranteed that only one decimal per string) and possibly with a '$' or 'USD'). Return a list of prices as floats. If there is a non-numeric entry (e.g. 'Out of Stock'), skip it.

*Hint*: 
- In class we checked the type of a variable as ```type(var)==int```. The function ```isinstance(var, [types])``` can check whether a variable is of several types.

```isinstance(12.43, [int, float])``` --> True

```python
...

```

```python
# Some tests
assert convert_to_numeric([12, '15.50', '$20', '30 USD', 'Out of Stock']) == [12.0, 15.5, 20.0, 30.0]
assert convert_to_numeric(['5', 7.25, '$8.00', 'Not available']) == [5.0, 7.25, 8.0]

```

```python
grader.check("q6")
```

### Submission

When you are ready to submit, uncomment (cmd+/ or ctrl+/) and run the cell below.

```python
# !git add *
# !git commit -m 'ps03 submitted'
# !git push -f
```

Tests for q1 (points: None)
```
>>> assert factorize(57) == [1, 3, 19, 57]
>>> assert factorize(59) == [1, 59]
>>> assert factorize(1) == [1], "Your function doesn't handle the 1 case correctly."
>>> assert factorize(-7) is None, 'Remember to check the number is positive.'
That's prime.
Please use a positive integer.
```


Tests for q2 (points: None)
```
>>> assert check_cyclic_equality([1, 2, 3], [1, 2, 3]) == 0
>>> assert check_cyclic_equality([1, 2, 3], [2, 3, 1]) == 2
>>> assert check_cyclic_equality([1, 2, 3], [1, 2]) == False
>>> assert check_cyclic_equality([], []) == 0
>>> assert check_cyclic_equality(['a', 'b', 'a'], ['b', 'a', 'a']) == 2
>>> assert check_cyclic_equality(['dog', 'cat', 'bird'], ['cat', 'bird', 'dog']) == 2
```


Tests for q3 (points: None)
```
>>> assert extract_vowels('aeiou') == [1, 1, 1, 1, 1, 5]
>>> assert extract_vowels('AEIOU') == [1, 1, 1, 1, 1, 5]
>>> assert extract_vowels('AaBbCcDdEeFfGgHhIiJjKkLlMmNnOoPpQqRrSsTtUuVvWwXxYyZz') == [2, 2, 2, 2, 2, 10]
>>> assert extract_vowels('bcdfg') == [0, 0, 0, 0, 0, 0]
>>> assert extract_vowels('Gettysburg College') == [1, 2, 1, 2, 1, 7]
>>> assert extract_vowels('') == [0, 0, 0, 0, 0, 0]
>>> assert extract_vowels('123!@#aei') == [1, 1, 1, 0, 0, 3]
```


Tests for q4 (points: None)
```
>>> file_list = ['hi.pdf', 'hello.doc', 'ps01.ipynb', 'goodbye.pdf', 'test.ipynb']
>>> assert find_filetype(file_list, '.pdf') == ['hi.pdf', 'goodbye.pdf']
>>> assert find_filetype(file_list, '.ipynb') == ['ps01.ipynb', 'test.ipynb']
>>> assert find_filetype(file_list, '.doc') == ['hello.doc']
>>> assert find_filetype(file_list, '.xlsx') == []
>>> assert find_filetype([], '.pdf') == []
```


Tests for q5 (points: None)
```
>>> assert remove_duplicates([1, 2, 2, 3, 1, 4]) == [1, 2, 3, 4]
>>> assert remove_duplicates([5, 6, 7]) == [5, 6, 7]
>>> assert remove_duplicates([9, 9, 9, 9]) == [9]
>>> assert remove_duplicates([]) == []
>>> assert remove_duplicates(['a', 'b', 'a', 'c', 'b']) == ['a', 'b', 'c']
>>> assert remove_duplicates('A man, a plan, a canal, Panama.') == ['A', ' ', 'm', 'a', 'n', ',', 'p', 'l', 'c', 'P', '.']
```


Tests for q6 (points: None)
```
>>> assert convert_to_numeric([10, '20', 30.5, '$40', '50 USD']) == [10.0, 20.0, 30.5, 40.0, 50.0]
>>> assert convert_to_numeric([1, 2.5, 3]) == [1.0, 2.5, 3.0]
>>> assert convert_to_numeric(['$5', '6 USD', '$7.50']) == [5.0, 6.0, 7.5]
>>> assert convert_to_numeric(['Out of Stock', 'N/A', 'Unavailable']) == []
>>> assert convert_to_numeric([100, '200', 'Not available', '$300', '400 USD', 'error']) == [100.0, 200.0, 300.0, 400.0]
>>> assert convert_to_numeric([]) == []
>>> assert convert_to_numeric(['1.99', '$2.50', '3.00 USD']) == [1.99, 2.5, 3.0]
```
